This tutorial was created using the wonderful [Google Python Developers](https://developers.google.com/edu/python/) course, and has been modified under the [Creative Common's licence 2.5](https://creativecommons.org/licenses/by/2.5/) 

Copyright 2010 Google Inc.
Licensed under the Apache License, Version 2.0
http://www.apache.org/licenses/LICENSE-2.0

Google's Python Class
http://code.google.com/edu/languages/google-python-class/

# Basic string exercises

## Honour Code

I **YOUR NAME, YOUR SURNAME**, confirm - by submitting this document - that the solutions in this notebook are a result of my own work and that I abide by the EDSA honour code (https://drive.google.com/file/d/1QDCjGZJ8-FmJE3bZdIQNwnJyQKPhHZBn/view?usp=sharing).  

Non-compliance with the honour code constitutes a material breach of contract.

## donuts
Given an int count of a number of donuts, return a string of the form `Number of donuts: <count>`, where `<count>` is the number passed in. However, if the count is 10 or more, then use the word 'many' instead of the actual count.<br>
So donuts(5)  returns `Number of donuts: 5` and <br>
   donuts(23) returns `Number of donuts: many`

<!-- grade_id: cell-a4742f59416f3b2b, points: 0, grade: false, solution: true -->
```python
def donuts(count):
    ### BEGIN SOLUTION
    if count >= 10:
        return 'Number of donuts: many'
    else:
        return 'Number of donuts: {}'.format(count) 
    ### END SOLUTION
    return None
```

Your code will be tested against the list of assertions below.  Make sure that it passes them.

<!-- grade_id: cell-f2a478f9944472ed, points: 1, grade: true, solution: false -->
```python
assert donuts(5) == 'Number of donuts: 5'
assert donuts(3) == 'Number of donuts: 3'
```

<!-- grade_id: cell-982cd7c38f927820, points: 1, grade: true, solution: false -->
```python
assert donuts(109) == 'Number of donuts: many'
assert donuts(501) == 'Number of donuts: many'
```

<!-- grade_id: cell-2c83810061b87661, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert donuts(10) == 'Number of donuts: many'
### END HIDDEN TESTS
```

```python
donuts(5)
```

```python
donuts(23)
```

Output:
```
'Number of donuts: many'
```

## both_ends
Given a string s, return a string made of the first 2 and the last 2 chars of the original string, so 'spring' yields 'spng'. However, if the string length is less than or equal to 2, return the empty string instead.

<!-- grade_id: cell-28f6442eeaf35f40, points: 0, grade: false, solution: true -->
```python
def both_ends(s):
    ### BEGIN SOLUTION
    if len(s) <= 2:
        return ''
    else:
        return s[0:2] + s[::-1][0:2][::-1]
    ### END SOLUTION
    return None
```

Your code will be tested against the list of assertions below.  Make sure that it passes them.

<!-- grade_id: cell-bf12251560fc0214, points: 1, grade: true, solution: false -->
```python
assert both_ends('spring') == 'spng'
assert both_ends('s') == ''
assert both_ends('hel') == 'heel'
assert both_ends('tomorrow') == 'toow'
```

<!-- grade_id: cell-807e7377b8d61c4d, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert both_ends('sp') == ''
### END HIDDEN TESTS
```

Output:
```
'spng'
```

```python
both_ends('spring')
```

```python
both_ends('sp')
```

Output:
```
''
```

## fix_start
Given a string s, return a string where all occurences of its first char have been changed to `*`, except do not change the first char itself.<br>
e.g. 'babble' yields `ba**le`<br>
Assume that the string is length 1 or more.<br>
Hint: s.replace(stra, strb) returns a version of string s where all instances of stra have been replaced by strb.

<!-- grade_id: cell-f8a217e1a99a57e5, points: 0, grade: false, solution: true -->
```python
def fix_start(s):
    ### BEGIN SOLUTION
    s1 = s.replace(s[0], '*')[1:]
    s2 = s1.replace(s1[0], s[0:1])[0]
    return s2 + s1
    ### END SOLUTION
    return None
```

Your code will be tested against the list of assertions below.  Make sure that it passes them.

<!-- grade_id: cell-187ab2dd46dbe486, points: 1, grade: true, solution: false -->
```python
assert fix_start('babble') == 'ba**le'
assert fix_start('turtle') == 'tur*le'
assert fix_start('kdswueknvnvmkloskk') == 'kdswue*nvnvm*los**'
assert fix_start('mmmmm') == 'm****'
```

<!-- grade_id: cell-54b27dd764c8cfcf, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert fix_start('hkasklkoowi') == 'hkasklkoowi'
### END HIDDEN TESTS
fix_start('babble')
```

Output:
```
'ba**le'
```

## MixUp
Given strings a and b, return a single string with a and b separated by a space `<a> <b>`, except swap the first 2 chars of each string.<br>
e.g.<br>
`'mix', pod' -> 'pox mid'`<br>
`'dog', 'dinner' -> 'dig donner'` <br>
Assume a and b are length 2 or more.

<!-- grade_id: cell-190c0b5a4d50dfda, points: 0, grade: false, solution: true -->
```python
def mix_up(a, b): 
    ### BEGIN SOLUTION
    first = b[0:2] + a[2:]
    second = a[0:2] + b[2:]
    return first + ' ' + second
    ### END SOLUTION
    return None
```

Your code will be tested against the list of assertions below.  Make sure that it passes them.

<!-- grade_id: cell-615dfa035a085ded, points: 2, grade: true, solution: false -->
```python
assert mix_up('dog', 'dinner') == 'dig donner'
assert mix_up('max', 'min') == 'mix man'
assert mix_up('today', 'tomorrow') == 'today tomorrow'
assert mix_up('somerandom', 'string') == 'stmerandom soring'
assert mix_up('hdjjsdfsd', 'klertpop') == 'kljjsdfsd hdertpop'
assert mix_up('ab', 'ba') == 'ba ab'
```

```python
mix_up('mix', 'pod')
```

Output:
```
'pox mid'
```

```python
mix_up('dog', 'dinner')
```

Output:
```
'dig donner'
```

## verbing
Given a string, if its length is at least 3, add 'ing' to its end.<br> Unless it already ends in 'ing', in which case add 'ly' instead. <br>If the string length is less than 3, leave it unchanged. <br>Return the resulting string.

<!-- grade_id: cell-9563b451db1937af, points: 0, grade: false, solution: true -->
```python
def verbing(s):
    ### BEGIN SOLUTION
    if len(s) < 3:
        return s
    elif (len(s) >= 3) & (s[::-1][0:3][::-1] == 'ing'):
        return s + 'ly'
    else:
        return s + 'ing'
    ### END SOLUTION
    return None
```

Your code will be tested against the list of assertions below.  Make sure that it passes them.

<!-- grade_id: cell-7900c6414257fbfa, points: 1, grade: true, solution: false -->
```python
assert verbing('dfjgkj') == 'dfjgkjing'
assert verbing('interestingly') == 'interestinglying'
```

<!-- grade_id: cell-e160ce7bbc45c572, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert verbing('interesting') == 'interestingly'
assert verbing('hello world') == 'hello worlding'
assert verbing('ing') == 'ingly'
### END HIDDEN TESTS
verbing('stop')
```

Output:
```
'stoping'
```

<!-- grade_id: cell-dca9868e623dd328, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert verbing('hj') == 'hj'
### END HIDDEN TESTS
verbing('stopping')
```

Output:
```
'stoppingly'
```

## not_bad
Given a string, find the first appearance of the substring 'not' and 'bad'.<br> If the 'bad' follows the 'not', replace the whole 'not'...'bad' substring with 'good'.<br>
Return the resulting string.

So 'This dinner is not that bad!' yields:
`This dinner is good!`

<!-- grade_id: cell-9320395523e2ade3, points: 0, grade: false, solution: true -->
```python
def not_bad(s):
    ### BEGIN SOLUTION
    not_i = s.find('not')
    bad_i = s.find('bad')
    try:
        if ('not' in s) & ('bad' in s) & (s.index('bad') > s.index('not')):
            return s.replace(s[s.find('not'): s.find('bad') + 3], 'good')
        else:
            return s
    except:
        return s
    ### END SOLUTION
    return None
```

Your code will be tested against the list of assertions below.  Make sure that it passes them.
IMPORTANT: Your code must not throw any errors, regardless of the string.

<!-- grade_id: cell-7dbbbf78c615b53e, points: 2, grade: true, solution: false -->
```python
assert not_bad('This dinner is not that bad') == 'This dinner is good'
assert not_bad('He is not a bad person') ==  'He is good person'
assert not_bad('not hello bad') == 'good'
assert not_bad('Being bad also mean being not good') == 'Being bad also mean being not good'
```

<!-- grade_id: cell-bbb2b7818f4ab45e, points: 1, grade: true, solution: false -->
```python
assert not_bad('hasd kfjeifnv dke') == 'hasd kfjeifnv dke'
```

```python
not_bad('')
```

Output:
```
''
```

```python
not_bad('This dinner is not that bad')
```

Output:
```
'This dinner is good'
```

## front_back
Consider dividing a string into two halves.<br>
If the length is even, the front and back halves are the same length.<br>
If the length is odd, we'll say that the extra char goes in the front half.<br>

e.g. `'abcde'`, the front half is `'abc'`, the back half `'de'`.

Given 2 strings, a and b, return a string of the form<br>
a-front + b-front + a-back + b-back

<!-- grade_id: cell-fff78aa159fefc12, points: 0, grade: false, solution: true -->
```python
def front_back(a, b):
    ### BEGIN SOLUTION
    a_no = int(len(a)/2)
    a_back = a[::-1][0:a_no][::-1]
    a_front = a[0:-a_no]
    b_no = int(len(b)/2)
    b_back = b[::-1][0:b_no][::-1]
    b_front = b[0:-b_no]
    return a_front + b_front + a_back + b_back
    ### END SOLUTION
    return None
```

<!-- grade_id: cell-b461f4ca751cdcb9, points: 3, grade: true, solution: false -->
```python
assert front_back('abcde', 'xyz',) == 'abcxydez'
assert front_back('binary', 'decimal',) == 'bindeciarymal'
assert front_back('dskjfkds', 'bcjfbdf',) == 'dskjbcjffkdsbdf'
assert front_back('odd', 'even',) == 'odevden'
assert front_back('front', 'back',) == 'frobantck'
```

```python
front_back('abcde', 'xyz')
```

Output:
```
'abcxydez'
```

```python
front_back('a', 'b')
```

Output:
```
''
```